```python
import pandas as pd
import numpy as np
import otter
grader = otter.Notebook()
```

# <mark style="background: #008030; color: #ffffff;" >B</mark> Week 3: Practice Testing & Debugging

------

### Instructions:

- Attempt & submit to pass the assessment
  - You can attempt any number of questions and in any order provided.
  - The more you practice, the better you'll get ;)
- Some of these questions can be done direclty in this notebook, while others should be attempted in a Unix Shell
- Be sute to check your work before submitting.
  - Do not remove the `'''` or other provided markings from the answer spaces.
  - Do not make any changes to this notebook outside of the answer spaces provided.
  
#### Submitting

- Reset your outputs before submitting
  - Select the `Kernel` menu, then either `Restart & Run Clear Output` or `Restart & Run All`
  - Don't forget to save your notebook after this step
- Submit your .ipynb file to Gradescope using Git & GitHub
- You can submit as many times as needed
- When reviewing results, ignore any results listed under "Public Tests"  
  (There are no "Public Tests" in this assignment)

For more information, see the assessment page.

#### Question 01 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(5 Points)

Consider a function `add()` that you want to use in your code. You cannot review the function's code, but know from its documentation that:
* You can invoke the function using `add(x, y)` where `x` and `y` are inputs.
* This function adds two integers together and returns the result.

Design a single unit test, using an assertion, to ensure the function does not also concatenate strings. If the test fails, your code should output "Unit Test failed".

```python
def question1_tests(add):
    assert add('x','y')=='xy','Unit Test failed'
    
    # Write your solution here

```

```python
grader.check("Question 01")
```

Output:
```
~/opt/anaconda3/lib/python3.9/site-packages/otter/check/notebook.py in check(self, question, global_env)
    166 
    167         else:
--> 168             test_path = self._resolve_nb_path(None)
    169             test_name = question
    170 

~/opt/anaconda3/lib/python3.9/site-packages/otter/check/notebook.py in _resolve_nb_path(self, nb_path)
    136         elif nb_path is None and glob("*.ipynb"):
    137             notebooks = glob("*.ipynb")
--> 138             assert len(notebooks) == 1, "nb_path not specified and > 1 notebook in working directory"
    139             nb_path = notebooks[0]
    140 

AssertionError: nb_path not specified and > 1 notebook in working directory
```

#### Question 02 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(5 Points)

The code below is supposed to print a festive message but has a bug (oh no!)

Find the bug using **static analysis**, then fix the code below.

```python
def saySentence(a, b, c, d, e, f, g):
   
    var7 =a + d + b + d + c + e + d + f + d + str(g) + e
    return var7
  

var1 = 'Happy'
var2 = 'New'
var3 = 'Year'
var4 = ' '
var5 = '!'
var6 = 'It\'s'
var7 = 2022

print(saySentence(var1, var2, var3, var4, var5, var6, var7))
```

Output:
```
Happy New Year! It's 2022!
```

#### Question 03 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(10 Points)

Consider a function `maximum()` that you want to use in your code. You cannot review the function's code, but know from its documentation that:
* You can invoke the function using `maximum(x, y)` where `x` and `y` are inputs.
* This function takes two integers and returns the higher of the two.
  * If x and y are the same, it that value.

Design unit tests, using assertions, to test the function's correctness. Find a case where the code returns the wrong value. The test that fails should output "Unit Test failed".

You can verify your tests by running the cell marked below.

```python
def question3_tests(maximum):
    assert maximum(1,2)==3,'Unit Test failed'
    # Write your solution here
    

```

```python
grader.check("Question 03")
```

Output:
```
~/opt/anaconda3/lib/python3.9/site-packages/otter/check/notebook.py in check(self, question, global_env)
    166 
    167         else:
--> 168             test_path = self._resolve_nb_path(None)
    169             test_name = question
    170 

~/opt/anaconda3/lib/python3.9/site-packages/otter/check/notebook.py in _resolve_nb_path(self, nb_path)
    136         elif nb_path is None and glob("*.ipynb"):
    137             notebooks = glob("*.ipynb")
--> 138             assert len(notebooks) == 1, "nb_path not specified and > 1 notebook in working directory"
    139             nb_path = notebooks[0]
    140 

AssertionError: nb_path not specified and > 1 notebook in working directory
```

#### Question 04 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(10 Points)

The code below is supposed to print every prime number between 2 and 20 but has a bug (oh no!)

Find the bug using **static analysis**, then fix the code below.

```python
from IPython.core import debugger
def print_primes():
    set_trace()

    i = 2
    j = 2
    isComposite = False
    result = ''

    while i <= 20:
        while j < i:
            if i % j == 0:
                isComposite = True
                break
            j = j + 1

        if not isComposite:
            result = result + str(i) + ' '

        iscomposite = False
        j = 2
        i = i + 1

    print(result)
    
print_primes()
```

Output:
```
2 3
```

#### Question 05 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(15 Points)

The following code has a bug.  
It is supposed to take string and repeat that string once for each character in that string, up to a maximum of 10 repeats.  
For example, the input 'foo' has 3 characters, so the output should be 'foofoofoo'.

Find the bug using **pdb**/**ipdb**, then fix the code below.

```python

def repeat_string(s):
    if type(s) != str:
        raise TypeError('This function only takes string inputs');
    
    count = 0
    n = len(s)
    s_text=""

    while count < len(s):
        s_text += s
        count += 1
        
        if count >= 10:
            break

    
    return s_text

s = 'foo'
print(repeat_string(s))
```

Output:
```
foofoofoo
```

#### Question 06 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(15 Points) 

The following code crashes when the `counter` variable becomes or exceeds 100. 

Use **pdb**/**ipdb** to find the value of `result` when the program crashed. You are recommended to do this using pdb/ipdb as a script from the terminal, as this sets up post mortem debugging automatically for you, otherwise there are less than 20 runs of this loop, making it easy to manually step through checking variables with smart breakpoint placement.

```python
result = 1
counter = 1
while counter < 100:
    result = result * 2
    counter = counter * 5 % result
raise Exception('Crashed!')
```

Find the value of `result` at the time of the crash, and enter it in the space provided

```python
result_value = 125
```

#### Question 07 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(20 Points)

Consider a function `remove()` that you want to use in your code. You cannot review the function's code, but know from its documentation that:
* You can invoke the function using `remove(x, y)` where `x` and `y` are inputs.
* This function takes two strings. All instances of the second string are removed from the first string. For example `remove('hello', 'l')` returns `heo`.

Design unit tests, using assertions, to test the function's correctness. Find a case where the code returns the wrong string. The test that fails should output "Unit Test failed".

You can verify your tests by running the cell marked below.

```python
def question7_tests(remove):
    try:
        remove('hello','l')
    except:
        print('sorry')
    # Write your solution here
    
```

```python
grader.check("Question 07")
```

Output:
```
~/opt/anaconda3/lib/python3.9/site-packages/otter/check/notebook.py in check(self, question, global_env)
    166 
    167         else:
--> 168             test_path = self._resolve_nb_path(None)
    169             test_name = question
    170 

~/opt/anaconda3/lib/python3.9/site-packages/otter/check/notebook.py in _resolve_nb_path(self, nb_path)
    136         elif nb_path is None and glob("*.ipynb"):
    137             notebooks = glob("*.ipynb")
--> 138             assert len(notebooks) == 1, "nb_path not specified and > 1 notebook in working directory"
    139             nb_path = notebooks[0]
    140 

AssertionError: nb_path not specified and > 1 notebook in working directory
```

#### Question 08 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(20 Points)

Consider a function `school_addition()` that you want to use in your code. You cannot review the function's code, but know from its documentation that:
* You can invoke the function using `school_addition(x, y)` where `x` and `y` are inputs.
* This function takes two integers between 0 and 99 (inclusive) and returns their sum.

Design unit tests, using assertions to test the function's correctness. Find a case where the code returns the wrong value. The test that fails should output "Unit Test failed".

```python
def question8_tests(school_addition):
    
    # Write your solution here
    
```

```python
grader.check("Question 08")
```

Tests for Question 01 (points: 5)
```
>>> def add_bad(x, y):
...     return x + y;
>>> 
>>> try:
...     out = '';
...     question1_tests(add_bad);
... except AssertionError:
...     out = 'One of your assertions has failed. Well done; your tests have exposed a bug.';
... except Exception as e:
...     raise e
... else:
...     out = 'All of your assertions have passed. Bug NOT found.';
...     
>>> out
'One of your assertions has failed. Well done; your tests have exposed a bug.'
```


Tests for Question 02 (points: 5)


Tests for Question 03 (points: 10)
```
>>> def maximum_bad(x, y):
...     if type(x) != int or type(y) != int:
...         raise TypeError('This function only takes integer inputs');
...     
...     if y > x:
...         return y;
...     elif x == y:               
...         return 0               
...     else:
...         return x;
>>> 
>>> try:
...     question3_tests(maximum_bad);
...     out = '';
... except AssertionError:
...     out = 'One of your assertions has failed. Well done; your tests have exposed a bug.';
... except Exception as e:
...     raise e
... else:
...     out = 'All of your assertions have passed. Bug NOT found.';
...     
>>> out
'One of your assertions has failed. Well done; your tests have exposed a bug.'
```


Tests for Question 04 (points: 10)
```
>>> print_primes()
2 3
```


Tests for Question 05 (points: 15)


Tests for Question 06 (points: 15)
```
>>> result_value
512
```


Tests for Question 07 (points: 20)
```
>>> def remove_bad(x, y):
...     if type(x) != str or type(y) != str:
...         raise TypeError('This function only takes string inputs');
...     result = x;
...     # if y is blank remove all characters (bug)
...     if y == '':
...         return '';
...     # otherwise perform code correctly
...     result.replace(y, '');
...     return result;
>>> 
>>> 
>>> try:
...     question7_tests(remove_bad);
...     out = '';
... except AssertionError:
...     out = 'One of your assertions has failed. Well done; your tests have exposed a bug.';
... except Exception as e:
...     raise e
... else:
...     out = 'All of your assertions have passed. Bug NOT found.';
...     
>>> out
'One of your assertions has failed. Well done; your tests have exposed a bug.'
```


Tests for Question 08 (points: 20)
```
>>> def school_addition_bad(x, y):
...     if type(x) != int or type(y) != int:
...         raise TypeError('This function only takes integer inputs');
...     if x < 0 or x > 100 or y < 0 or y > 100:
...         raise ValueError('Inputs must be in the range 0 to 99 (inclusive)');
...     
...     x_1 = x // 10;
...     x_2 = x % 10;
...     y_1 = y // 10;
...     y_2 = y % 10;
...     
...     r_2 = x_2 + y_2;
...     c = r_2 // 10;
...     r_2 = r_2 % 10;
...     
...     r_1 = x_1 + y_1 + c;
...     c = r_1 // 10;
...     r_1 = r_1 % 10;
...     
...     r_0 = c;
...     r = r_0 * 10*2 + r_1 * 10**1 + r_2;
...     
...     return r;
>>> 
>>> 
>>> try:
...     question8_tests(school_addition_bad);
...     out = '';
... except AssertionError:
...     out = 'One of your assertions has failed. Well done; your tests have exposed a bug.';
... except Exception as e:
...     raise e
... else:
...     out = 'All of your assertions have passed. Bug NOT found.';
...     
>>> out
'One of your assertions has failed. Well done; your tests have exposed a bug.'
```
